The Neuropeptide Sulfakinin, a peripheral regulator of insect behavioral switch between mating and foraging

Author response:

The following is the authors’ response to the previous reviews

Joint Public Review:

Summary:

The behavioral switch between foraging and mating is important for resource allocation in insects. This study characterizes the role of sulfakinin and the sulfakinin receptor 1 in changes in olfactory responses associated with foraging versus mating behavior in the oriental fruit fly (Bactrocera dorsalis), a significant agricultural pest. This pathway regulates food consumption and mating receptivity in other species; here the authors use genetic disruption of sulfakinin and sulfakinin receptor 1 to provide strong evidence that changes in sulfakinin signaling modulate antennal responses to food versus pheromonal cues and alter the expression of ORs that detect relevant stimuli.

Strengths:

The authors utilize multiple complementary approaches including CRISPR/Cas9 mutagenesis, behavioral characterization, electroantennograms, RNA sequencing and heterologous expression to convincingly demonstrate the involvement of the sulfakinin pathway in the switch between foraging and mating behaviors. The use of both sulfakinin peptide and receptor mutants is a strength of the study and implicates specific signaling actors.

Weaknesses:

The authors demonstrate that SKR is expressed in olfactory neurons, however there are additional potential sites of action that may contribute to these results.

Recommendations for the authors:

The authors have addressed most of the issues raised by the reviewers. Below are a few outstanding issues.

(1) Lines 68-69 describe "control of B. dorsalis include the use of the behavioral responses to semiochemicals" but does not describe what these responses are or how behavior is modulated.

The sentence was revised as “Control of B. dorsalis include the use of the reproductive and feeding behavioral responses to semiochemicals” (lines 69 in the revision).

(2) Statistical analysis for 9 hour starved females at 5 minutes is missing in Figure 1D and S1.

We had added statistical analysis for 9 hour starved females at 5 minutes in the revised Figures 1D and S1, respectively (lines 578).

(3) The legend in Figure S2 should be revised as it is not clear from the figure which of the odors are food associated odors.

As suggested, we added food odor label in the revised Figure S2 (lines 666).

(4) Line 167: "Therefore, the upregulated OR genes in starved WT flies, OR7a.4, OR7a.8 and OR10a, were activated by the pheromonal components, while down regulated genes, OR49a and OR63a, were activated by food volatiles." Based on the data, this sentence is incorrect - Therefore, the upregulated OR genes in starved WT flies, OR7a.4, OR7a.8 and OR10a, were activated by the food components, whereas downregulated genes, OR49a and OR63a, were activated by pheromonal components."

We are sorry for our mistake. We had corrected it (lines 168-169).

(5) Line 192: "The coordinated action of sulfakinin on mutiple downstreams,..." should be revised to "downstream pathways or tissues" or simply removing "multiple downstream".

As suggested, we removed “multiple downstream”. See line 192.

(6) Reference formatting is inconsistent: see line 207 vs line 208.

We had corrected it as “(Wu et al., 2019)” (lines 207).

(7) Lines 241-244 The broad discussion regarding the evolution and ancestral function of CCK here and the phylogeny in Figure S6 are peripheral to the authors claims.

As suggested, we removed the section and the Figure S6 in the revision.

Changes from the earlier version to the revised version

Abstract

Abstract

Behavioral strategies for foraging and reproduction in the oriental fruit fly (Bactrocera dorsalis) are alternative options for resource allocation and are controlled by neuropeptides. Here we show that the behavioral switch between foraging and reproduction is associated with changes in antennal sensitivity. Starved flies became more sensitive to food odors while suppressing their response to opposite-sex pheromones. The gene encoding sulfakinin receptor 1 (SkR1) was significantly upregulated in the antennae of starved flies, so we tested the behavioral phenotypes of null mutants for the genes encoding the receptor (skr1-/-) and its ligand sulfakinin (sk-/-). In both knockout lines, the antennal responses shifted to mating mode even when flies were starved. This suggests that sulfakinin signaling via SkR1 promotes foraging while suppressing mating. Further analysis of the mutant flies revealed that sets of odorant receptor (OR) genes were differentially expressed. Functional characterization of the differentially expressed ORs suggested that sulfakinin directly suppresses the expression of ORs that respond to opposite-sex hormones while enhancing the expression of ORs that detect food volatiles. We conclude that sulfakinin signaling via SkR1, modulating OR expressions and leading to altered antenna sensitivities, is an important component in starvation-dependent behavioral change.

Introduction

Earlier version: Foraging and mating behaviors are alternative options for resource allocation in insects. Starved and reproductively immature individuals intensify their foraging activity, whereas satiated and reproductively mature individuals increase their mating activity to facilitate reproduction (Chen et al., 2023; Cheriyamkunnel et al., 2021; Ebrahim et al., 2023; Ren et al., 2021). Such behavioral changes involve complex neural pathways triggered by neuropeptide signaling (González Segarra et al., 2023; Lin et al., 2022; Nässel and Zandawala, 2022; Schoofs et al., 2017; Sun et al., 2023; Zhang et al., 2022). The key regulators have been identified in the fruit fly Drosophila melanogaster, including the neuropeptides sulfakinin (Sk) (Wang et al., 2022; Wu et al., 2019) and short neuropeptide F (sNPF) (Ko et al., 2015; Root et al., 2011). Tyramine and diuretic hormone 31 then resolve the conflict between foraging and mating behavior in the central nervous system (Cheriyamkunnel et al., 2021; Lin et al., 2022).

Revised version: Foraging and mating behaviors are alternative options for resource allocation in insects. Starved and reproductively immature individuals intensify their foraging activity, whereas satiated and reproductively mature individuals increase their mating activity to facilitate reproduction (Chen et al., 2023; Cheriyamkunnel et al., 2021; Ebrahim et al., 2023; Ren et al., 2021). Such behavioral changes involve complex neural pathways triggered by neuropeptide signaling (González Segarra et al., 2023; Nässel and Zandawala, 2022; Schoofs et al., 2017; Sun et al., 2023; Zhang et al., 2022a; Zhang et al., 2022b). The key regulators have been identified in the fruit fly Drosophila melanogaster, including the neuropeptides sulfakinin (Sk) (Wang et al., 2022; Wu et al., 2019) and short neuropeptide F (sNPF) (Ko et al., 2015; Root et al., 2011). Tyramine and diuretic hormone 31 then resolve the conflict between foraging and mating behavior in the central nervous system (Cheriyamkunnel et al., 2021; Zhang et al., 2022a).

In mammals, cholecystokinin (CCK) transmits satiety by activating cholecystokinin receptor 1, leading to the expression of genes encoding other anorectic and orexigenic peptides and receptors in vagal neurons (Beutler et al., 2017; Burdyga et al., 2008; Rachel et al., 2017). CCK is also involved in the central integration of sensory inputs to regulate sexual behavior (Bloch et al., 1987; Micevych et al., 2004). In insects, Sk, the ortholog of CCK, and its signals via two receptors, SkR1 (CCKLR-17D3 of D. melanogaster) and SkR2 (CCKLR-17D1), regulate sub-behavioral steps during foraging, food consumption, and mating in D. melanogaster (Fedina et al., 2023; Wang et al., 2022). For example, Sk suppresses the expression of the Gr64f gustatory receptor gene, which inhibits feeding as a signal for the onset of satiety (Guo et al., 2021). Sk also inhibits both male and female sexual behavior (Wu et al., 2019). However, the complexity of Sk pathways, spatial and temporal dynamics of multiple ligands and receptors, is exemplified by cell-specific signaling via SkR1 for the promotion of mating receptivity in virgin females (Wang et al., 2022), which is an opposite action of the inhibition of sexual behavior (Wu et al., 2019). Modulation of peripheral sensory inputs, including odorant perception, is involved in the control of mating and foraging behavior in addition to the central neuronal circuits.

Earlier version: The oriental fruit fly, Bactrocera dorsalis, is a destructive agricultural pest insect. Control of B. dorsalis includes the use of the behavioral responses to semiochemicals, such as the components in the sex pheromones and in the food smell (Jaffar et al., 2023). Chemical components in the sex pheromone of this species includes 2,3,5-trimethyl pyrazine, 2,3,5,6-tetramethylpyrazine, ethyl laurate, ethyl myristate, ethyl cis-9-Hexadecenoate, and ethyl palmitate (Chen et al., 2023; Ren et al., 2021) and in the food smells are ethyl benzoate, ethyl butyrate, and methyl eugenol, among others. (Kamala Jayanthi et al., 2014; Zhu et al., 2018). The physiology and molecular mechanisms underlying the behavioral modulation involved in the different types of chemical attractants have not yet been studied in this species of insect.

Revised version: The oriental fruit fly, Bactrocera dorsalis, is a destructive agricultural pest insect. Control of B. dorsalis includes the use of the reproductive and feeding behavioral responses to semiochemicals, such as the components in the sex pheromones and in the food smell (Jaffar et al., 2023). Chemical components in the sex pheromone of this species includes 2,3,5-trimethyl pyrazine, 2,3,5,6-tetramethylpyrazine, ethyl laurate, ethyl myristate, ethyl cis-9-Hexadecenoate, and ethyl palmitate (Chen et al., 2023; Ren et al., 2021) and in the food smells are ethyl benzoate, ethyl butyrate, and methyl eugenol, among others. (Kamala Jayanthi et al., 2014; Zhu et al., 2018). The physiology and molecular mechanisms underlying the behavioral modulation involved in the different types of chemical attractants have not yet been studied in this species of insect.

We hypothesized that neuropeptide Sk-SkR and other neuropeptidergic systems are involved in the switch between foraging success and mating behaviors in the B. dorsalis. We were specifically interested in the alteration in odorant perceptions in the antennae leading to the success in foraging and/or mating depending on the physiological status. Sequential approaches, starting with the unbiased test for involvement of odorant perception in the behavioral switch, led us to further in-depth investigations of the role of neuropeptide and Sk-SkR systems. We demonstrate that sulfakinin signaling via SkR1 in the antennae alters the expression of sets of odorant receptor (OR) genes during starvation enhancing the foraging success rate. Our results thus provide insight into sulfakinin signal on peripheral olfactory system that is coordinated with the central neuronal circuits for modulation of behavior.

Results

Altered antennal sensitivity is associated with starvation-induced foraging and satiation-induced mating success

We developed an assay to investigate the behavioral switch between foraging and mating induced by the physiological state of B. dorsalis (Fig. 1A, B). Foraging success increased with the duration of the starvation period (Fig. 1D, and S1), whereas courtship and copulation rates decreased, also with the starvation length-dependent manner (Fig. 1G, H). Starved flies tended to move directionally toward the food source located in the center of the arena (Fig. 1C).

Starvation alters the behavior and olfactory responses of B. dorsalis.

(A) Schematic illustration of the experimental design. (B) Schematic illustration of the mating behavior assay device. (C) Representative foraging trajectories in the 100 mm diameter arenas within a 15-min observation period of flies starved for different durations. (D) Cumulative successful foraging within a 15-min observation period of male flies starved for different durations. Data are means ± SEM, n = 60 flies for each condition. Different letters above the error bars indicate significant differences (one-way ANOVA followed by Tukey’s multiple comparisons test; P < 0.05). (E‒F) EAG responses to artificial food in fed and starved male (E) and female (F) flies (unpaired t-test). (Top) Representative EAG recordings. (G) Cumulative courtship rate within a 15-min observation period of flies starved for different durations (n = 78, 76, 78, 84 and 91, respectively, from fed to 12 h). (H) Cumulative copulation rate within a 15-min observation period of the flies starved for different durations (n = 78, 76, 79, 83 and 93, respectively, from fed to 12 h). (G–H) Kruskal-Wallis and post hoc Mann-Whitney U tests were applied. (I‒J) EAG responses to body extracts in fed and starved male (I) and female (J) flies. (Top) Representative EAG recordings. Data are means ± SEM, n = 10–12 antennae per genotype (unpaired t-test; *P < 0.05, **P < 0.01, ***P < 0.001).

We therefore tested antennal sensitivities toward artificial food and opposite-sex body extracts using an electroantennogram (EAG). Flies of both sexes starved for 12 h showed a greater EAG response to food odor (Fig. 1E, F) but a weaker response to opposite-sex body extracts (Fig. 1I, J). Furthermore, starved flies also showed greater EAG responses to food volatile components and weaker responses to sex pheromone components (Fig. S2). This suggested antenna-mediated olfaction is associated with and required for successful foraging and mating.

Starvation-induced SkR1 mRNA in the antennae is required for the switch from mating to foraging behavior

To identify signaling molecules responsible for the olfactory-dependent behavioral switch, we compared the expression profiles of neuropeptide signaling components in the antennae of fed and starved flies. We found that the components of three neuropeptide signaling systems were significantly upregulated in the antennae of starved flies, namely short neuropeptide F receptor (sNPFR), SIFamideR1, and SkR1, the latter showing the strongest induction (Fig. 2A). Sk did not show the expression change in the antennae, but significantly increased in the head in comparison between those in starvation and after re-feeding, suggesting that the circulatory Sk released from the central nervous system could be the source for activation of SkR1 in the antennae (Fig. S3).

Removed in the revised version: We introduced a 7-bp deletion into the sulfakinin gene (Sk) using the CRISPR/Cas9 system, causing a frameshift that abolished the sequence region encoding two predicted mature peptides: Sk-1 and Sk-2 (Fig. 2B, S4A, B). We confirmed the absence of an anti-sulfakinin immunoreactive signal in the central nervous system of the sk−/− flies (Fig. S4C). Only SkR1 was upregulated in the starved flies among two predicted Sk receptor genes SkR1 and SkR2 (Fig. 2A, and S6). We therefore generated a SkR1 null mutant (skr1−/−) by introducing a 187-bp deletion using CRISPR/Cas9, resulting in a frameshift that resulted in a premature stop codon (Fig. 2B, and S5A−C). We confirmed the absence of the anti-SkR1 immunoreactive signal in the antennae of the skr1−/− flies (Fig. S5D).

Removed in the revised version: Starved sk−/− and skr1−/− flies showed the same locomotor activity and body size as wild-type (WT) controls when we measured their velocity and body length (Fig. 2C, and S7). However, both null mutants consumed more food than WT flies (Fig. 2D), consistent with previous results in D. melanogaster (Guo et al., 2021). Interestingly, sk−/− and skr1−/− male flies consumed more food at all times compared to the wild type. However, the corresponding mutant females consumed more food only at night. This may be due to the high basal feeding rate of females during the daytime, masking the increase in feeding in the knockout of Sk signaling. As shown in Figure 2D, WT females ate more food than males, and WT females consumed more food during the daytime. Another major difference between the null mutants and WT controls was that the mutant flies were less successful at foraging even after starvation, similar to the characteristics of satiated WT flies (Fig. 3A, B). However, foraging was not completely abolished. In addition, starved mutant females (Fig. 3C) and males (Fig. 3D) showed high mating success rates, as same rates as fed WT flies. We tested the behavioral phenotypes of heterozygous mutant of Sk knockout flies. The results showed that foraging and mating behaviors of Sk heterozygous mutants were unaffected during starvation, suggesting the mutants are completely recessive (Fig. S8).

Starvation-induced expression of neuropeptide system components and phenotypes of null mutant flies.

(A) Changes in the expression of neuropeptide system components in the antennae of starved flies. The relative expression values are the fold-changes compared to the control. Data are means ± SEM, n = 3 (unpaired t-test; *P < 0.05, **P < 0.01). (B) Altered amino acid sequence and mature peptides in the null mutants sk-/- and skr1-/- compared to wild-type (WT) flies. (C) Mean velocity of flies representing each genotype during a 15-min observation period. Data are means ± SEM (Kruskal–Wallis test). (D) Knockout of sulfakinin and SkR1 increases food consumption in B. dorsalis at different times of day. Each experiment consisted of one fly, and at least 15 flies were contained for the assay. Different lower-case letters indicate significant differences between treatments (one-way ANOVA followed by Tukey’s multiple comparisons test; P < 0.05).

Added in the revised version: We introduced a 7-bp deletion into the sulfakinin gene (Sk) using the CRISPR/Cas9 system, causing a frameshift that abolished the sequence region encoding two predicted mature peptides: Sk-1 and Sk-2 (Fig. 2B, S4A, B). We confirmed the absence of an anti-sulfakinin immunoreactive signal in the central nervous system of the sk-/- flies (Fig. S4C). Only SkR1 was upregulated in the starved flies among two predicted Sk receptor genes SkR1 and SkR2 (Fig. 2A). We therefore generated a SkR1 null mutant (skr1-/-) by introducing a 187-bp deletion using CRISPR/Cas9, resulting in a frameshift that resulted in a premature stop codon (Fig. 2B, and S5A−C). We confirmed the absence of the anti-SkR1 immunoreactive signal in the antennae of the skr1-/- flies (Fig. S5D).

Added in the revised version: Starved sk-/- and skr1-/- flies showed the same locomotor activity and body size as wild-type (WT) controls when we measured their velocity and body length (Fig. 2C, and S6). However, both null mutants consumed more food than WT flies (Fig. 2D), consistent with previous results in D. melanogaster (Guo et al., 2021). Interestingly, sk-/- and skr1-/- male flies consumed more food at all times compared to the wild type. However, the corresponding mutant females consumed more food only at night. This may be due to the high basal feeding rate of females during the daytime, masking the increase in feeding in the knockout of Sk signaling. As shown in Figure 2D, WT females ate more food than males, and WT females consumed more food during the daytime. Another major difference between the null mutants and WT controls was that the mutant flies were less successful at foraging even after starvation, similar to the characteristics of satiated WT flies (Fig. 3A, B). However, foraging was not completely abolished. In addition, starved mutant females (Fig. 3C) and males (Fig. 3D) showed high mating success rates, as same rates as fed WT flies. We tested the behavioral phenotypes of heterozygous mutant of Sk knockout flies. The results showed that foraging and mating behaviors of Sk heterozygous mutants were unaffected during starvation, suggesting the mutants are completely recessive (Fig. S7).

Sk-SkR1 signaling is required for changes in behavior and olfactory sensitivity induced by starvation.

(A‒B) Cumulative successful foraging within a 15-min observation period of fed and starved flies with different genotypes. The successful foraging of females (A) and males (B) was measured separately. Data are means ± SEM, n = 60 flies for each condition. Different letters above the error bars indicated significant differences (one-way ANOVA followed by Tukey’s multiple comparisons test; P < 0.05). (C) Cumulative copulation rates within a 15-min observation period of fed and starved flies with different genotypes (n = 92, 86, 88, 106, 80 and 76, respectively, from fed WT to starved skr1-/-). (D) Cumulative courtship rates within a 15-min observation period of fed and starved flies with different genotypes (n = 91, 85, 86, 98, 80 and 77, respectively, from fed WT to starved skr1-/-). (C–D) Kruskal-Wallis and post hoc Mann-Whitney U tests were applied (*P < 0.05, **P < 0.01, ***P < 0.001). (E‒F) EAG responses to artificial food in fed and starved flies of different genotypes. The EAG responses to artificial food in females (E) and males (F) were measured separately. (Top) Representative EAG recordings. (G‒H) EAG responses to body extracts in fed and starved flies of different genotypes. The EAG responses to body extracts of females (G) and males (H) were measured separately. (Top) Representative EAG recordings. Data are means ± SEM, n = 10–12 antennae for each condition. Different letters above the error bars indicated significant differences (one-way ANOVA followed by Tukey’s multiple comparisons test; P < 0.05).

Earlier version: We next tested antennal sensitivity to food odors and body extracts by EAG in each genotype using both fed and starved flies in each case. The starved sk−/− and skr1−/− flies showed significantly weaker responses to food odors compared to the starved WT flies, although the response of sk−/− flies was slightly stronger than that of fed WT flies (Fig. 3E, F and S8), which was similar to the behavioral response for a successful foraging rate (Fig. 3A, B). The antennae of starved sk−/− and skr1−/− flies showed the same EAG response to opposite-sex body extracts as fed WT flies (Fig. 3G, H and S9), which was significantly higher than that of starved WT flies.

Revised version: We next tested antennal sensitivity to food odors and body extracts by EAG in each genotype using both fed and starved flies in each case. The starved sk-/- and skr1-/- flies showed significantly weaker responses to food odors compared to the starved WT flies, although the response of sk-/- flies was slightly stronger than that of fed WT flies (Fig. 3E, F and S8), which was similar to the behavioral response for a successful foraging rate (Fig. 3A, B). The antennae of starved sk-/- and skr1-/- flies showed the same EAG response to opposite-sex body extracts as fed WT flies (Fig. 3G, H and S8), which was significantly higher than that of starved WT flies.

Sets of ORs are differentially expressed when comparing satiated/starved WT and sk-/- flies

To gain insight into the molecular basis of alteration of olfactory sensitivities by the Sk-SkR1 system, we extracted RNA from the head epidermal tissues (including antennae) of fed and starved flies for sequencing (RNA-Seq). We compared starved WT and starved sk-/- flies with satiated WT flies, focusing on the differentially expressed OR genes (Fig. 4A). We identified three OR genes (OR7a.4, OR7a.8 and OR10a) that were upregulated in the starved WT flies but unchanged in starved sk-/-flies compared to satiated WT controls, and two (OR49a and OR63a) that were downregulated in starved WT flies but unchanged in starved sk-/-flies (Fig. 4B). OR67d.3 and OR74a mRNA levels were similar in the satiated and starved WT flies. The differentially expressed genes identified by RNA-Seq were further validated by real-time reverse-transcription PCR (qRT-PCR) (Fig. 4B).

The transition of olfactory responses is associated with the expression of different sets of ORs in OR neurons induced by Sk-SkR1 signaling.

(A) Principal component analysis (PCA) using differentially expressed genes obtained from pairwise comparisons between different treatments. (B) The expression profiles of the candidate OR genes in wild-type (WT) and sk-/- flies determined by qRT-PCR. The OR names with red fonts represent the set of ORs for food odor and blue fronts represent the set of ORs for sex pheromone. Data are mean relative expression levels ± SEM. R1−R3 represent biological replicates. Different lower-case letters indicate significant differences between treatments (one-way ANOVA followed by Tukey’s multiple comparisons test; P < 0.05). (C) Quantification of calcium levels following the response of candidate ORs to sex pheromones and food odors at a concentration of 10-4 M (n = 3‒6). Different lower-case letters indicate significant differences between treatments (one-way ANOVA followed by Tukey’s multiple comparisons test; P < 0.05). (D) Dose–response curves of candidate ORs to sex pheromone components such as ethyl laurate (EL), ethyl palmitate (EP), 2,3,5-trimethyl pyrazine (TMP), and 2,3,5,6-tetramethylpyrazine (TTMP). (E) Dose–response curves of candidate ORs to food odorants such as ethyl benzoate (EBE), ethyl butyrate (EBU), diethyl maleate (DM), and methyl eugenol (ME). In both cases, n = 5. (F) Co-localization of Orco (red) and SkR1 (green) neurons in B. dorsalis antennae. Scale bars = 50 μm.

Sulfakinin regulates the expressions of starvation-induced ORs that detect food volatiles and starvation-suppressed ORs that detect pheromones

Earlier version: The OR genes identified in the RNA-Seq experiments were expressed in HEK 293 cells followed by the evaluation of functional responses to odorants using calcium mobilization assays. Tests with single high doses of volatiles representing food or sex pheromones allowed the functional categorization of ORs based on odorant specificities (Fig. 4C). OR10a was activated by ethyl benzoate (EBE), OR7a.8 was activated by ethyl butyrate (EBU) or diethyl maleate (DM), and OR7a.4 was activated by methyl eugenol (ME). The responses to these food odor volatiles were all dose-dependent, with OR10a showing the strongest response to EBE (Fig. 4E). In contrast, OR63a and OR49a showed dose-dependent responses to sex pheromone components but not to food odor volatiles (Fig. 4C, D). Therefore, the upregulated OR genes in starved WT flies, OR7a.4, OR7a.8 and OR10a, were activated by the pheromonal components, while down reulated genes, OR49a and OR63a, were activated by food volatiles. Finally, we determined whether SkR1 was expressed in OR neurons by testing the antennae for immunoreactivity to SkR1 and odorant receptor co-receptor (Orco) antisera. A subset of clustered OR neurons showed immunoreactivity to SkR1 (Fig. 4F), indicating that sulfakinin acts directly on SkR1-positive OR neurons.

Revised version: The OR genes identified in the RNA-Seq experiments were expressed in HEK 293 cells followed by the evaluation of functional responses to odorants using calcium mobilization assays. Tests with single high doses of volatiles representing food or sex pheromones allowed the functional categorization of ORs based on odorant specificities (Fig. 4C). OR10a was activated by ethyl benzoate (EBE), OR7a.8 was activated by ethyl butyrate (EBU) or diethyl maleate (DM), and OR7a.4 was activated by methyl eugenol (ME). The responses to these food odor volatiles were all dose-dependent, with OR10a showing the strongest response to EBE (Fig. 4E). In contrast, OR63a and OR49a showed dose-dependent responses to sex pheromone components but not to food odor volatiles (Fig. 4C, D). Therefore, the upregulated OR genes in starved WT flies, OR7a.4, OR7a.8 and OR10a, were activated by the food volatiles, while down regulated genes, OR49a and OR63a, were activated by pheromonal components. Finally, we determined whether SkR1 was expressed in OR neurons by testing the antennae for immunoreactivity to SkR1 and odorant receptor co-receptor (Orco) antisera. A subset of clustered OR neurons showed immunoreactivity to SkR1 (Fig. 4F), indicating that sulfakinin acts directly on SkR1-positive OR neurons.

Discussion

Dynamic changes in the olfactory sensitivity of insect antennae are associated with behavioral transitions in insects. Sensory inputs via olfactory perception often trigger certain behaviors, including foraging and mating (Anton and Rössler, 2021; Gadenne et al., 2016). Fine-tuning of the antennal OR repertoire would be an efficient way to co-opt behavioral outputs (Anton and Rössler, 2021; Xu et al., 2024) as we have shown here for the oriental fruit fly B. dorsalis. Changes in odorant sensitivity during transitions between foraging and mating have been reported in previous studies. For example, the antennae of female D. suzukii become significantly more sensitive to the host fruit odor isoamyl acetate after mating (Crava et al., 2019). Similarly, starved Aedes aegypti mosquitoes downregulate the sensitivity of lactic acid receptor neurons in grooved peg sensilla, which are used during host search (Davis, 1984; Oliveira et al., 2017), whereas satiated mosquitoes increase the sensitivity of OR neurons in trichoid sensilla to compounds representing the oviposition site (Siju et al., 2010). In this study, we determined the mechanisms underlying the changes in antennal sensitivity when switching between foraging and mating behavior. Sk-SkR1 signaling was required to increase the success of foraging and to suppress mating during a period of starvation. The activation of SkR1 on the antennal OR neurons upregulates the expression of ORs that detect food volatiles while suppressing those that detect pheromones.

Earlier version: The coordinated actions of sulfakinin on multiple downstreams, the antennae and other parts of the peripheral and central nervous system in B. dorsalis, are not yet fully understood. Sulfakinin is generally known as a satiety signal that reduces food intake and suppresses the digestive system in D. melanogaster and a number of different invertebrate species (Nässel and Wu, 2022), whereas we found that the Sk-SkR1 signal is required to promote foraging by making the antennae more sensitive to food odorants during starvation. In D. melanogaster, sulfakinin inhibits the sugar receptor GR64 in the proboscis and proleg tarsi, thus reducing food ingestion (Guo et al., 2021). Our data also support the role of sulfakinin as a satiety signal because feeding rate was increased in sk−/− and skr1−/− null mutants (Fig. 2D). Sulfakinin was shown to inhibit both male and female sexual behavior in D. melanogaster (Wu et al., 2019), in agreement with our results. In addition to the previous study showing the roles of central nervous system in Sk-mediated inhibition of sexual behavior in D. melanogaster, we described a Sk direct action on the peripheral olfactory system in B. dorsalis. The underlying mechanism is the secretion of the mature peptide Dsk-2 (but not Dsk-1) by Dsk-MP1 and Dsk-MP3 neurons, which then triggers downstream neurons expressing SkR1 in the central nervous system to suppress male mating behavior (11). However, there was a contradictory report showing the role of sulfakinin promoting the receptivity in mating in virgin females (Wang et al., 2022). The activity of sulfakinin peptides Sk1 and Sk2, through the receptors SkR1 and SkR2, is therefore likely to be complex, with spatially and temporally dynamic effects in multiple downstream systems.

Revised version: The coordinated actions of sulfakinin on the antennae and other parts of the peripheral and central nervous system in B. dorsalis, are not yet fully understood. Sulfakinin is generally known as a satiety signal that reduces food intake and suppresses the digestive system in D. melanogaster and a number of different invertebrate species (Nässel and Wu, 2022), whereas we found that the Sk-SkR1 signal is required to promote foraging by making the antennae more sensitive to food odorants during starvation. In D. melanogaster, sulfakinin inhibits the sugar receptor GR64 in the proboscis and proleg tarsi, thus reducing food ingestion (Guo et al., 2021). Our data also support the role of sulfakinin as a satiety signal because feeding rate was increased in sk-/- and skr1-/- null mutants (Fig. 2D). Sulfakinin was shown to inhibit both male and female sexual behavior in D. melanogaster (Wu et al., 2019), in agreement with our results. In addition to the previous study showing the roles of central nervous system in Sk-mediated inhibition of sexual behavior in D. melanogaster, we described a Sk direct action on the peripheral olfactory system in B. dorsalis. The underlying mechanism is the secretion of the mature peptide Dsk-2 (but not Dsk-1) by Dsk-MP1 and Dsk-MP3 neurons, which then triggers downstream neurons expressing SkR1 in the central nervous system to suppress male mating behavior (Wu et al., 2019). However, there was a contradictory report showing the role of sulfakinin promoting the receptivity in mating in virgin females (Wang et al., 2022). The activity of sulfakinin peptides Sk1 and Sk2, through the receptors SkR1 and SkR2, is therefore likely to be complex, with spatially and temporally dynamic effects in multiple downstream systems.

Although successful foraging and EAG responses to food odors were strongly reduced in both null mutants, there was still a greater response than observed in fed WT controls (Fig. 3A, B, E, F). This incomplete suppression suggests that other mechanisms, in addition to the main Sk-SkR1 signal, are relevant to the response. Indeed, sNPF and tachykinin signaling have both been described in the modulation of olfactory sensitivity in D. melanogaster, with sNPF enhancing sensitivity to food odors and tachykinin inhibiting sensitivity to aversive odors (Ignell et al., 2009; Ko et al., 2015; Root et al., 2011). Furthermore, neuropeptide F (NPF) and SIFamide also regulate foraging behavior at the level of the olfactory circuit (Jennifer and Yi, 2013; Martelli et al., 2017; Wang et al., 2013). Notably, NPF and tachykinin also regulate mating behavior at the peripheral level. Tachykinin inhibits courtship behavior by mediating the perception of an anti-aphrodisiac pheromone in male D. melanogaster. (Shankar et al., 2015). NPF neurons are also required to detect female sex pheromones (Gendron et al., 2014; Kuo et al., 2012). Similarly, SIFamide also regulates mating behavior, although it is not yet clear whether it acts via the olfactory circuit (Sellami and Veenstra, 2015; Terhzaz et al., 2007). We have not yet investigated the roles of other neuropeptides in B. dorsalis foraging and mating behavior, but we found that starvation-induced the neuropeptide receptor genes SkR1, sNPFR, and SIFamideR1 (P < 0.05). In addition, tachykinin, SIFamide, and ion transport peptide (ITP)-like peptide were also induced by starvation, although the increase was not statistically significant. Therefore, other neuropeptides in addition to sulfakinin are likely to be involved in the switch between different sets of ORs in the antenna as a means to control behavior.

Removed in the revised version: The evolutionary origin of CCK/sulfakinin signaling can be traced back to the common ancestor of the Bilateria (Mirabeau and Joly, 2013; Nässel and Wu, 2022; Staljanssens et al., 2011; Tinoco et al., 2021; Yu and Smagghe, 2014). Most coelomate and pseudocoelomate phyla produce CCK/sulfakinin peptides and orthologs of the corresponding receptors. Sulfakinin is involved in foraging, satiety, locomotion, development, aggression and reproductive behavior in insects (Fedina et al., 2023; Nässel and Wu, 2022; Nassel and Zandawala, 2019) although some insect lineages have lost both the peptide and its receptors (Consortium, 2010; Hauser et al., 2010). CCK is also involved in satiety, memory, anxiety, and depression in mammals (Jens et al., 2007; Lau et al., 2023; Rana et al., 2022; Shen et al., 2019). The many functions of CCK/sulfakinin in diverse taxonomic groups frequently include signaling between the digestive system (endoderm) and the neuronal system (ectoderm) (Nässel and Wu, 2022). This suggests an ancestral function associated with coelomic development, and that additional complexity has evolved in each lineage.

In summary, we found that the neuropeptide sulfakinin mediates the transition between foraging and mating behavior by activating the receptor SkR1 and reprogramming the peripheral OR repertoire. We propose a working model to characterize this mechanism, in which Sk-SkR1 regulates the expression of different sets of ORs, leading to the behavioral switch (Fig. 5). The modulation of peripheral sensory systems by sulfakinin, coordinating with the central neural circuit also controlled by sulfakinin, is the major trigger for a successful behavioral switch. We have thus provided new insight into the control of behavior by sulfakinin in B. dorsalis, which is one of the most widely distributed, fecund, and invasive insect pests to threaten global agriculture (Wan and Yang, 2016; Zhao et al., 2024).

Schematic representation of peripheral olfactory remodeling in OR neurons of the antenna that arbitrate between mating and foraging behavior in B. dorsalis.

Starvation increases the abundance of SkR1 in the OR neurons (red arrows) and SkR1 signaling induces the expression of genes encoding ORs that sense food odors, resulting in successful foraging. Satiation suppresses SkR1 expression and induces the expression of genes encoding ORs that sense opposite-sex pheromones, leading to successful mating (blue arrows). Dashed arrows represent the additional possible pathways that have not been tested in this study, but not excluded in the model. Please see the discussion for details of additional possible factors modulating odorant sensitivity relevant to satiety.

Materials and methods

Insects

The wild-type (WT) oriental fruit flies were originally collected from Hainan province, China in 2008 and were reared in the laboratory (27.5 ± 0.5 °C, 75 ± 5% relative humidity, 14-h photoperiod). Larvae and adults were fed with artificial food as previously described (Shen et al., 2013). Artificial food for adults, consisting of yeast powder, honey, sugar, ascorbic acid and water, was also used as a food source for foraging behavior assays.

Behavioral assays

Virgin females and males, 3-day old, were placed each in groups of 25 flies and reared at 27.5 °C and 60% humidity, and 9-day old flies were subjected to behavioral assays. Food consumption rates of individual flies were measured by the amount Bromophenol Blue dye (BPB) in the fly after the food including 0.1% (BPB) was provided for 1 h for ad libitum feeding, which is long enough time to allow satiation to occur. BPB was measured from decapitated and homogenized bodies in 1 ml PBS buffer with 1% Triton X-100. Absorbance of the supernatant after centrifugation (13000 rpm) for 3 min was measured at 614 nm on a xMarkerTM Microplate Spectrophotometer (Bio-Rad). Flies were starved for 12 h prior to the experiments with only water. Foraging behavior assay was performed according to previously published methods with modifications (Li et al., 2022). In brief, single flies were introduced into chambers (100 mm diameter x 10 mm height) and acclimatized for 20 min. The foraging behavior was recorded using a camera (Sony HDR CX405) for 15 min at 30 fps. Successful foraging was counted when the fly stayed 5 s or longer within the 5 mm radius of the center where 10 μL of artificial food was placed. The experiments consisted of the assays of 3 biological replications for each experimental condition. The behavior was analyzed by EasyFlyTracker (Qu et al., 2022) with the default parameters.

Mating behavior assay was conducted in the courtship chambers (diameter: 3.5 cm; height: 1.5cm per layer). The virgin females and males were cold anesthetized at ‒20 °C for 1 min and introduced into the courtship chamber separated by a removable transparent film. After 1 h acclimation, the film was removed to allow the paired flies to get in contact, and mating behavior was recorded by a video camera for 15 min. We measured the level of courtship in mutant males and receptivity in mutant females respectively. For courtship assay, a mutant male (sk-/- or skr-/-) and a WT virgin female were introduced into the chambers. Courtship rate is the percentage of the mutant males performs attempted copulation, which was used to measure mutant courtship level. For female receptivity assay, a mutant female (sk-/-or skr-/-) and a WT virgin male were introduced into the chambers. Copulation rate is the percentage of the mutant females had engaged in copulation, which is used to measure the mutant female receptivity. The number of flies that had performed attempted copulation or engaged in copulation by the end of each 1 min interval were summed within 15 min and plotted as a percentage of total flies for each time point.

Electroantennograms

Electroantennograms were used to record the stimulus-dependent changes in the electrical potential summed over the whole antenna. To test volatiles were food odor or body extract of each sex on the flies conditioned for satiated and starved. Different genotypes of the flies were wild-type, Sk mutant, and SkR1 mutants. Female body extracts were tested by male antennae and male body extracts by female antennae (Table S2). Body extracts were prepared by washing 15 9-day old virgin flies of each sex in 1 ml of hexane for 4 hours (Chen et al., 2023; Dweck et al., 2015; Ebrahim et al., 2023) at room temperature with gentle shaking.

For electroantennograms, similar to previous study (Li et al., 2022), the head (including the antennae) of fly was excised and a glass capillary filled with 0.9% NaCl was inserted into the cut end as the reference electrode. The tip of an antenna was placed into another glass capillary filled with 0.9% NaCl as the recording electrode. The reference electrode was connected to a grounding electrode, while the recording electrode was connected to a 10× high-impedance DC amplifier (Syntech; The Netherlands). The signal was then sent to an analog/digital converter (IDAC-4, USB, Syntech) and transferred to a computer. Finally, the data were analyzed and saved by using GC-EAD software (Syntech). A filter paper strip (3 × 0.75 cm), loaded with 10 μL of the odor solution or body extract, was placed into a pipette and delivered a stimulus into a constant, humidified airstream flowing at 100 ml/min by a stimulus controller (Syntech, Hilversum, Netherlands). Two stimuli were made for 0.5s duration with a 30s interval. Average EAG response, as shown by mV changes, was used for the calculation. EAG for sex pheromones and body surface extracts were measured from 21:00‒23:00, and food odor was measured from 9:00-11:00.

Quantitative real-time reverse transcription PCR

Antennal tissues from flies of different genotypes with different treatments were dissected and immediately stored at ‒80 °C. Total RNA was extracted with Trizol Reagent (Invitrogen) following the manufacturer’s protocol. First-strand cDNA was synthesized using the PrimeScript 1st Strand cDNA Synthesis Kit (Takara, Dalian, China) and specific primers designed using the online software, Primer 3 (http://primer3.ut.ee/, Table S1). Quantitative real-time reverse transcription PCR (qRT-PCR) was performed using a CFX Connect Real-Time System (Bio-Rad, Hercules, CA) with NovoStart SYBR PCR SuperMix kit (Novoprotein, Shanghai, China). Two reference genes (Shen et al., 2010), α-tubulin (GenBank: GU269902) and RPS3 (GenBank: XM_011212815), were used selected to normalize the expression levels for the calculation in qBASE (Hellemans et al., 2007). Each experiment comprised 3 independent biological replicates and two technical replicates.

CRISPR/Cas9-mediated mutants

In vitro synthesis of single guide RNA was performed as previously described (Li et al., 2022). In brief, the suitable target sites of Sk and SkR1 genes were searched by CasOT (Xiao et al., 2014), and the corresponding single guide RNAs (sgRNAs) were synthesized in vitro and purified using the GeneArt Precision gRNA Synthesis Kit (Invitrogen, Thermo Fisher Scientific, Waltham, MA, USA). We designed and synthesized one sgRNA for the Sk gene (Fig. S4) and two single guide RNAs for the SkR1 gene (Fig. S5).

The sgRNA and cas9 protein (Invitrogen) were gently pipetted and mixed in nuclease-free water to form a mixture of 400 ng/μL sgRNA and 300 ng/μL Cas9 protein. Eggs collected within 20 min of oviposition were injected within 2 h with the mixture of sgRNA and Cas9 protein using an IM300 Electric Microinjector (Narishige, Tokyo Japan). The injected eggs were incubated at 27.5 ± 0.5 °C and 75 ± 5% relative humidity for hatching.

Mutant screening and off-target analysis

Heterozygous G1 offspring were obtained using 9-day-old virgin adults (G0) crossed to wild-type opposite sex. Suitable mutant G1 (sk+/- and skr1+/-) were backcrossed to WT flies for at least eight generations, then homozygous mutant flies (sk-/- and skr1-/-) were obtained by self-crossing. Genomic DNA was extracted from individual flies (G0) or a single hind leg (G1 and later generations) by incubating tissue samples in 30 μL InstaGene Matrix (Bio-Rad, Hercules, CA, USA) at 56 °C for 30 min and inactivating the enzyme at 100 °C for 10 min. Fragments of 300–500 bp spanning the sgRNA target sites were amplified using the primers (Table S1). The fragments were genotyped using the QIAxcel Advanced System (Qiagen, Hilden Germany), and potential edited products were directly sequenced or cloned into vector pGEM-T Easy (Promega, Wisconsin, Madison, USA) for Sanger sequencing to identify the mutation.

Earlier version: To assess off-target effects, CasOT was used to search potential off-target sites for the Sk and SkR1 genes. All loci with one potential mismatch allowed in the seed region and up to three potential mismatches in the non-seed region were searched on the genome sequence (https://db.cngb.org/search/project/CNP0003192/). Two potential off-target sites were found for Sk gene and three for SkR1 gene. PCR amplification and sequencing of fragments including potential off-target sites were performed to check for mutations. In examination of the potential off-target sites with > 4 bp mismatches to the sgRNA sequences, no mutations were found in the off-sites (Fig. S10).

Revised version: To assess off-target effects, CasOT was used to search potential off-target sites for the Sk and SkR1 genes. All loci with one potential mismatch allowed in the seed region and up to three potential mismatches in the non-seed region were searched on the genome sequence (https://db.cngb.org/search/project/CNP0003192/). Two potential off-target sites were found for Sk gene and three for SkR1 gene. PCR amplification and sequencing of fragments including potential off-target sites were performed to check for mutations. In examination of the potential off-target sites with > 4 bp mismatches to the sgRNA sequences, no mutations were found in the off-sites (Fig. S9).

Immunohistochemistry

Rabbit anti-BdSkR1 antibody was generated by immunization with the peptide N’-RQGLRKCRDQWP-C’. Mouse anti-BdOrco antibody was raised by immunization with two peptides, N’-GTNPNGLTRKQE-C’ and N’-HWYDGSEEAKTF-C’. Peptide were synthesized and conjugated with keyhole limpet hemocyanin through additional cysteine residues at their N termini by Zoonbio Biotechnology (Nanjing, China). Rabbit anti-BdSk antibody was a gift from Dr. Chuan Zhou (Wu et al., 2020). The antisera used to recognize Sk peptide were raised in New Zealand white rabbits using the synthetic peptide N’-GGDDQFDDYGHMRFG-C.

Immunohistochemistry was performed as previously described (Li et al., 2022). In brief, whole brains of 9-day-old flies were dissected in phosphate-buffered saline (PBS), fixed overnight in 4% paraformaldehyde (PFA) in PBS at 4 °C, and washed (3 × 15 min) in PBST (0.5% Triton X-100 in PBS). The tissues were blocked in 5% normal goat serum for 20 min at room temperature, then incubated in primary antibody for 48 h at 4 °C. After three washes in PBST (3 × 15 min), the tissues were then incubated in secondary antibody overnight at 4 °C. After washing in PBST again, the tissues were mounted on a slide with 100% glycerol for imaging. Adult antennae were dissected and embedded in Tissue-Tek O.C.T. Compound (Sakura, Tokyo, Japan). Sections (18 μm) were prepared using a Thermo HM525 NX microtome at ‒22 °C and then fixed in 4% PFA in PBS at room temperature for 2 h. The subsequent procedures were as described above for the brain tissues. Primary antibody used: rabbit anti-Sk (1:1000), rabbit anti-SkR1 (1:500), mouse anti-Orco (1:200). Secondary antibodies used: goat anti-rabbit IgG conjugated to Alexa Fluor 488 (1:1000, Cell Signaling Technology), goat anti-mouse IgG conjugated to Alex Fluor 647 (1:1000, Cell Signaling Technology). Samples were imaged using an LSM780 confocal microscope (Carl Zeiss, Jena, Germany) and processed with Image J software (Schneider et al., 2012).

RNA-seq analysis

Total RNA was extracted from the head epidermal tissues, including the antennae, of fed WT flies, starved WT flies, and starved sk-/- flies using TRIzol reagent (Invitrogen). Three replicates were performed. Library construction and sequencing was performed by Novogene with Illumina HiSeq2000 platform (Novogene Bioinformatics Technology Co.Ltd, Beijing, China). The raw data were analyzed after filtering the low-quality sequences by fastp with the default parameters (Chen et al., 2018). Sequences were aligned to the B. dorsalis genome (https://db.cngb.org/search/project/CNP0003192/) using Hisat2 (version 2.2.1) (Kim et al., 2015; Yang et al., 2023). Samtools (version 1.7) was used to convert sam files to bam files in the process (Li et al., 2009). Finally, to calculate the count, we use featureCounts (version 2.0.1) (Liao et al., 2014). The expression level of genes from the RNA sequencing was normalized by the FPKM method. Differential expression analysis was performed using the DESeq2 R package (1.16.1) (Anders and Huber, 2010).

Functional heterologous expression of ORs and calcium imaging

Heterologous expressions of ORs were made in Human Embryonic Kidney-293 (HEK-293) cells, which were cultured in Dulbecco’s modified Eagle medium (DMEM, 11995065, Gibco, Thermo Fisher Scientific, USA) containing 10% Fatal Bovine Serum (FBS, 10091148, Gibco) and 1% penicillin-streptomycin-Amphotericin B solution (C0224-100 ml, Beyotime) with 5% CO2 at 37 °C. Cells for transfection were seeded onto a 96-well black well plate (165305, Thermo Fisher Scientific, USA). Approximately 50%–70% confluent HEK-293 cells were transfected by 2 μg pcDNA3.1 plasmid containing respective olfactory receptors with BdOrco of B. dorsalis using TransIT®-LT1 Transfection Reagent (MIR 2300, Mirusbio) for 24 h.

Earlier version: For calcium imaging experiment, the culture media was replaced to 100 μl of 2.5 μM Fluo-4 AM (F14201, Invitrogen) and incubated at 37 °C with 5% CO2 for 1 h. Subsequently, after another 30 min of incubation at room temperature, the culture media containing Fluo-4 was removed, and the cells were washed three times with Hanks’ balanced salt solution (HBSS; 14065056, Gibco, Thermo Fisher Scientific, USA) in order to fully remove the Fluo-4 residue. Then, 99 ul of HBSS was added to each well and plate was placed in the dark at RT to be tested. Finally, 1 μL of test odor was added to each well and fluorescence immediately was measured for at least 80 s. Finally, images were captured by a Zeiss Observer 7 inverted microscope (Carl Zeiss, Jena, Germany). Test odors were dissolved in dimethylsulfoxide (DMSO, ST038-100ml, Beyotime) and prepared five final concentrations: 10−7, 10−6, 10−5,

Revised version: For calcium imaging experiment, the culture media was replaced to 100 μl of 2.5 μM Fluo-4 AM (F14201, Invitrogen) and incubated at 37 °C with 5% CO2 for 1 h. Subsequently, after another 30 min of incubation at room temperature, the culture media containing Fluo-4 was removed, and the cells were washed three times with Hanks’ balanced salt solution (HBSS; 14065056, Gibco, Thermo Fisher Scientific, USA) in order to fully remove the Fluo-4 residue. Then, 99 ul of HBSS was added to each well and plate was placed in the dark at RT to be tested. Finally, 1 μL of test odor was added to each well and fluorescence immediately was measured for at least 80 s. Finally, images were captured by a Zeiss Observer 7 inverted microscope (Carl Zeiss, Jena, Germany). Test odors were dissolved in dimethylsulfoxide (DMSO, ST038-100ml, Beyotime) and prepared five final concentrations: 10−7, 10−6, 10−5, 10−4 M and 10−3 M. the fluorescent ratio (ΔF/F) was calculated as peak intensity divided by basal intensity (intensity before ligand perfusion). For all experiment, the response to DMSO was used as negative control.

Removed in the revised version: 10-4‒M and 10-3‒M. the fluorescent ratio (ΔF/F) was calculated as peak intensity divided by basal intensity (intensity before ligand perfusion). For all experiment, the response to DMSO was used as negative control.

Statistical analysis

For qRT-PCR result analysis, relative expression levels were calculated using qBase (Hellemans et al., 2007) based on the expression of the reference genes. Statistical analysis is performed using GraphPad Prism and indicated in each figure legend. Data were first verified for normal distribution using the Shapiro-Wilk normality test. One-way ANOVA, followed by Tukey’s multiple comparisons (if normal distribution) and Kruskal-Wallis and post-hoc Mann-Whitney U-tests (if non-normal distribution), was used for comparisons among multiple groups. Values are reported as means ± standard errors of the mean (SEM).